## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## KEYWORDS('inner product spaces','complex matrices','unitary diagonalization')

## DBsubject('Linear Algebra')
## DBchapter('Inner Product Spaces')
## DBsection('Unitary Diagonalization')
## Date('2015/03/28')
## Author('J. Sylvestre')
## Institution('University of Alberta, Augustana Campus')
## TitleText1('')
## EditionText1('')
## AuthorText1('')
## Section1('')
## Problem1('')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGmatrixmacros.pl",  # for mbox, display_matrix_mm
  "PGchoicemacros.pl",
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context("Complex");
Context()->parens->set("<" => {type => "Vector", close => ">"});

$n = 3;

$t = random(2,8);
foreach $i (1..4) {
	$b[$i] = random(-1,1,2);
	$a[$i] = $t * $b[$i];
}

$eig_shift = random(-3,3);
$eig[1] = Compute("1 - 2*$t + $eig_shift");
$eig[2] = Compute("1 + $eig_shift");
$eig[3] = Compute("1 + 2*$t + $eig_shift");

$H = Matrix([
	[1, Complex($a[1],$a[2]), 0],
	[Complex($a[1],-$a[2]), 1, Complex($a[3],$a[4])],
	[0, Complex($a[3],-$a[4]), 1]
]) + $eig_shift * Value::Matrix->I($n);


# Unnormalized eigenvecs
# eig1
$p[1] = ColumnVector(
  Compute("($b[1]+$b[2]*i)*($b[3]+$b[4]*i)"),
  Compute("-2*($b[3]+$b[4]*i)"),
  Compute("2"),
);
# eig2
$p[2] = ColumnVector(
  Compute("-2"),
  Compute("0"),
  ($p[1]->value)[0]->conj,
);
# eig3
$p[3] = ColumnVector(
  ($p[1]->value)[0],
  (-1)*( ($p[1]->value)[1] ),
  ($p[1]->value)[2],
);

# normalized eigenvecs
foreach $i (1..3) {
	$pconj[$i] = ColumnVector( map(conj($_),$p[$i]->value) );
	$pnormsq[$i] = $pconj[$i] . $p[$i];
	$q[$i] = $p[$i] / Compute("sqrt($pnormsq[$i])");
}

Context()->texStrings;
BEGIN_TEXT
$BBOLD $BITALIC Recall: $EITALIC $EBOLD Hermitian matrices always have real eigenvalues.
$PAR
The Hermitian matrix \(H=$H\) has a three distinct, real eigenvalues \(\lambda_1 < \lambda_2 < \lambda_3\).
$PAR
\{mbox([
	'\(\lambda_1 =\)', ans_rule(10),
	'\(\lambda_2 =\)', ans_rule(10),
	'\(\lambda_3 =\)', ans_rule(10),
]) \}
END_TEXT
Context()->normalStrings;
ANS($eig[1]->cmp);
ANS($eig[2]->cmp);
ANS($eig[3]->cmp);

Context()->texStrings;
BEGIN_TEXT
$PAR
To achieve \( U^{\ast} HU =
  \{ display_matrix_mm([['\lambda_1','',''],['','\lambda_2',''],['','','\lambda_3']]) \}
\)
we could use \(U =
  \{display_matrix_mm([['\mathbf{u}_1','\mathbf{u}_2','\mathbf{u}_3']])\}
\),
\{mbox(
  'where \(\mathbf{u}_1 = \)',
  $q[1]->ans_array(20),
  ', \(\mathbf{u}_2 = \)',
  $q[2]->ans_array(20),
  ' and \(\mathbf{u}_3 = \)',
  $q[3]->ans_array(20),
  '.'
)\}
END_TEXT
Context()->normalStrings;

sub chk_vec {
	my ($student,$i) = @_;

	# check parallel by projecting student answer onto $eigvec
	# if parallel this projection should be equal to the original student answer
	my $student_proj_diff = $pnormsq[$i] * $student - ($pconj[$i] . $student) * $p[$i];
	return 0 if (! $student_proj_diff->isZero);

	# check unit vector
	my $student_conj = ColumnVector( map(conj($_),$student->value) );
	my $student_normsq_diff = $student_conj . $student - Complex(1,0);
	return ($student_normsq_diff->isZero);
}

ANS($q[1]->cmp(
	showCoordinateHints => 0,
	checker=> sub {
		my ( $correct, $student, $ansHash ) = @_;
		return chk_vec($student,1);
	}
));
ANS($q[2]->cmp(
	showCoordinateHints => 0,
	checker=> sub {
		my ( $correct, $student, $ansHash ) = @_;
		return chk_vec($student,2);
	}
));
ANS($q[3]->cmp(
	showCoordinateHints => 0,
	checker=> sub {
		my ( $correct, $student, $ansHash ) = @_;
		return chk_vec($student,3);
	}
));

ENDDOCUMENT();       # This should be the last executable line in the problem.
